#############################################
###       Generated from PreTeXt source      
###       on 2026-02-23T21:53:00-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(A.4.10)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "scaffold.pl",
  "contextForm.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context("Numeric")->variables->add(c=>'Real',m=>'Real',n=>'Real',q=>'Real',r=>'Real');
Context()->reduction->set('(-x)-y'=>0,'(-x)+y'=>0);
$giv1=Formula("53m");
$ans1=Formula("m*53");
$ev1=$ans1->cmp(checker=>sub{
  my ($correct,$student,$ans)=@_;
  Value->Error("This is the original expression") if ($student eq $giv1);
  Value->Error("This is algebraically equivalent to the given expression, but using the commutative property of multiplication would give you something else") if (($student == $correct) && ($student ne $correct));
  return ($student eq $correct);
});
Context()->flags->set(reduceConstants=>0);
Context()->lists->set(AbsoluteValue=>{open=>"(",close=>")"});
Context()->parens->set("("=>{removable=>0,type=>"AbsoluteValue",close=>")"});
$giv2=Formula("3*(5 n)");
$ans2=Formula("(3*5)*n");
$alt2=Formula("15*n");
$ev2=$ans2->cmp(checker=>sub{
  my ($correct,$student,$ans)=@_;
  $student->{test_at}=[[0,0,-1,0,0,0],[0,0,1,0,0,0]];
  $correct->{test_at}=[[0,0,-1,0,0,0],[0,0,1,0,0,0]];
  Value->Error("This is algebraically equivalent to the given expression, but using the associative property of multiplication would give you something else") if (($student == $giv2) || (($student == $correct) && ($student ne $correct) && ($student ne $alt2)));
  return (($student eq $correct) or ($student eq $alt2));
});
$giv3=Formula("q+84");
$ans3=Formula("84+q");
$ev3=$ans3->cmp(checker=>sub{
  my ($correct,$student,$ans)=@_;
  Value->Error("This is the original expression") if ($student eq $giv3);
  Value->Error("This is algebraically equivalent to the given expression, but using the commutative property of addition would give you something else") if (($student == $correct) && ($student ne $correct));
  return ($student eq $correct);
});
$giv4=Compute("x+(20+c)");
$amb4=Compute("x+20+c");
$ans4=Compute("(x+20)+c");
$ev4=$ans4->cmp(checker=>sub{
  my ($correct,$student,$ans)=@_;
  for my $x ($giv4,$student,$amb4,$correct) {
    $x->{test_at}=[[-21,0,0,0,0,-21],[-21,0,0,0,0,-19],[-19,0,0,0,0,-21],[-19,0,0,0,0,-19]];
  }
  Value->Error("This expression is algebraically equivalent to the given expression, but using the associative property of addition would give you something else") if (($student == $giv4) || ($student == $amb4) || (($student == $correct) && ($student ne $correct)));
  return ($student eq $correct);
});
$giv5=Compute("3*(r+7)")->reduce;
Context("Form")->variables->are(r=>'Real');
$ans5=Compute("3*r+3*7");

############################################################
# Body
############################################################

############################################################
# Scaffold
############################################################
Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Use the commutative property of multiplication to write an equivalent expression to [`[$giv1]\text{.}`]

[_]{$ev1}{10}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
To use the commutative property of multiplication, change the order the two factors are multiplied:

[```
                53m=m\cdot 53
              \text{.}```]

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Use the associative property of multiplication to write an equivalent expression to [`[$giv2]\text{.}`]

[_]{$ev2}{10}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
To use the associative property of multiplication, leave factors written in their original order, but change the grouping symbols so that a different multiplication has higher priority:

[```
                3(5n)=(3\cdot5)n
              \text{.}```]

You may further simplify by multiplying the two numbers:

[```\newcommand{\amp}{&}
\begin{aligned}
3(5n)\amp=(3\cdot5)n\\
\amp=15n
\end{aligned}
```]

but that is going a step beyond simply using the associative property.

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Use the commutative property of addition to write an equivalent expression to [`[$giv3]\text{.}`]

[_]{$ev3}{10}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
To use the commutative property of addition, change the order the two terms are added:

[```
                q+84=84+q
              \text{.}```]

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Use the associative property of addition to write an equivalent expression to [`[$giv4]\text{.}`]

[_]{$ev4}{10}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
To use the associative property of addition, leave terms written in their original order, but change the grouping symbols so that a different addition has higher priority:

[```
                x+(20+c)=(x+20)+c
              ```]

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Use the distributive property to write an equivalent expression to [`[$giv5]`] that has no grouping symbols.

[_]{$ans5}{10}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
To use the distributive property, multiply the number outside the parentheses, [`3\text{,}`] with each term inside the parentheses:

[```\newcommand{\amp}{&}
\begin{aligned}
3(r+7)\amp=3\cdot r+3\cdot7\\
\amp=3r+21\text{.}
\end{aligned}
```]

END_PGML_SOLUTION

Section::End();

Scaffold::End();

############################################################
# End Problem
############################################################

ENDDOCUMENT();
