#############################################
###       Generated from PreTeXt source      
###       on 2026-02-23T21:53:00-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(A.5.4.31)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGchoicemacros.pl",
  "contextArbitraryString.pl",
  "contextFraction.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context("Fraction");
@d = (2..9)[NchooseK(8,2)];
($a, $b) = map{-Fraction(random_coprime([1..2*$_], [$_]))} (@d);
Context("ArbitraryString");
$operator = ($a > $b) ? '>' : '<';
$answer = Compute($operator)->cmp(
    checker => sub {
        my ($correct,$student,$ans) = @_;
        $correct = $correct->value;
        $student = $student->value;
        if ($student =~ /^~~s*$correct~~s*$/) {return 1}
        elsif ($student !~ /^~~s*>~~s*$/ && $student !~ /^~~s*<~~s*$/ && $student !~ /^~~s*=~~s*$/)
            {Value->Error("Answer with >, <, or =.")}
        else {return 0};
    }
);

############################################################
# Body
############################################################

BEGIN_PGML
Decide if one given number is greater than, less than, or equal to another given number.

[`[$a]`] [_]{$answer}{3} [`[$b]`]

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
