#############################################
###       Generated from PreTeXt source      
###       on 2026-02-23T21:52:58-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.10.14)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "scaffold.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "parserNumberWithUnits.pl",
  "parserLinearRelation.pl",
  "parserGraphTool.pl",
  "Pd.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context("LinearRelation")->variables->are(t=>'Real',y=>'Real');
@eq=(LinearRelation("y=16-0.8t"),LinearRelation("y=2+0.6t"));
$m=MultiAnswer(@eq)->with(singleResult=>0,
checker=>sub{
my($c,$s,$self)=@_;
my @s=@{$s};
my @c=@{$c};
if (($c[0]==$s[0]&&$c[1]==$s[1])||($c[0]==$s[1]&&$c[1]==$s[0])){return [1,1]}
elsif ($c[0]==$s[0]||$c[1]==$s[0]){return [1,0]}
elsif ($c[0]==$s[1]||$c[1]==$s[1]){return [0,1]}
else {return [0,0]}
});
$gt=GraphTool("{line,solid,(0,16),(5,12)}","{line,solid,(0,2),(5,5)}")->with(
availableTools=>["LineTool"],showInStatic=>0,bBox=>[-1,20,15,-1],
ticksDistanceX=>5,ticksDistanceY=>5,minorTicksX=>0,minorTicksY=>0,xAxisLabel=>'t',
cmpOptions=>{list_checker=>sub{
my($c,$s,$a,$v)=@_;
return 0 if $ans->{isPreview};
my $score=0; my @errors; my $count=1;
my ($x1,$y1)=$c->[0]->extract(3)->value;
my ($x2,$y2)=$c->[0]->extract(4)->value;
my ($x3,$y3)=$c->[1]->extract(3)->value;
my ($x4,$y4)=$c->[1]->extract(4)->value;
my @std=($y1-$y2,$x2-$x1,$x1*$y2-$x2*$y1,$y3-$y4,$x4-$x3,$x3*$y4-$x4*$y3);
my $p=sub{my ($A,$B,$C,$pt)=@_; my ($x,$y)=$pt->value; return $A*$x+$B*$y+$C==0;};
my @l=(0,0);
for (@$s) {
my $nth=Value::List->NameForNumber($count++);
if ($p->(@std[0..2],$_->extract(3))&&$p->(@std[0..2],$_->extract(4))) {$l[0]=1;$score += 1; next}
if ($p->(@std[3..5],$_->extract(3))&&$p->(@std[3..5],$_->extract(4))) {$l[1]=1;$score += 1; next}
push(@errors, "The $nth line graphed is incorrect.");
}
return (($l[0]&&$l[1]) ? $score : 0,@errors);
}});
$years=NumberWithUnits("10 yr");
$population=Real(8);

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      xmin=0,
      xmax=15,
      ymin=0,
      ymax=20,
      xtick={0,5,10,15},
      ytick={0,5,...,20},
      minor xtick={0,1,...,15},
      minor ytick={0,1,...,20},
      xlabel={\(t\)},
      width=0.47\orccaprintwidth,
      axis line style={->}
    ]
    \addplot+[-{Kite},domain=0:15]{16-0.8*x} node[pos=0.25,sloped,below] {\(y=16-0.8t\)};
    \addplot+[-{Kite},domain=0:15]{2+0.6*x} node[pos=1,sloped,above left] {\(y=2+0.6t\)};
    \addplot[guideline,->] coordinates {(0,16) (5,16) (5,12)};
    \addplot[guideline,->] coordinates {(0,2) (5,2) (5,5)};
    \addplot[soliddot,color=firstcolor] coordinates {(0,16) (5,12)};
    \addplot[soliddot,color=secondcolor] coordinates {(0,2) (5,5)};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
A college has a north campus and a south campus. The north campus has [`16`] thousand students, but enrollment has been declining by [`0.8`] thousand students per year. The newer south campus has only [`2`] thousand students, but enrollment has been increasing by [`0.6`] thousand students per year. If these trends continue, in how many years would the two campuses have the same number of students?

END_PGML

############################################################
# Scaffold
############################################################
Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Write two equations that form a system for this scenario.

[@KeyboardInstructions('Use [|t|]* as the variable for how much time has passed, and [|y|]* for the variable representing how many students are at one campus.')@]**

[_]{$m}{16}

[_]{$m}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The north campus has initial population [`16`] and decreases with rate [`0.8`] thousand students per year. So one equation is [`y=16-0.8t\text{.}`]

The south campus has initial population [`2`] and increases with rate [`0.6`] thousand students per year. So the other equation is [`y=2+0.6t\text{.}`]

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Plot the two lines.

[_]{$gt}{1}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
For plotting, it may be helpful to write the system with the decimals converted to fractions:

[```\newcommand{\amp}{&}
              \left\{
              \begin{alignedat}{3}
              y \amp {}={} \amp -\frac{4}{5}t \amp {}+{} \amp 16 \\
              y \amp {}={} \amp \frac{3}{5}t \amp {}+{} \amp 2
              \end{alignedat}
              \right.
              ```]

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a Cartesian grid with two intersecting lines; one line passes through (0,16) with slope -0.8; the other line passes through (0,2) with slope 0.6; the lines cross at (10,8)")@]*  <<

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Based on the graph, how long will it be until the two campuses have the same number of students?

[_]{$years}{16}

How many students will each campus have at that time?

[_]{$population}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The lines cross at [`(10,8)\text{.}`] So after [`10`] years, each campus will have [`8`] thousand students.

END_PGML_SOLUTION

Section::End();

Scaffold::End();

############################################################
# End Problem
############################################################

ENDDOCUMENT();
