#############################################
###       Generated from PreTeXt source      
###       on 2026-02-20T13:16:09-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.11.5.27)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "PCCmacros.pl",
  "PGchoicemacros.pl",
  "parserImplicitPlane.pl",
  "contextForm.pl",
  "contextFraction.pl",
  "parserGraphTool.pl",
  "Pd.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
# Get first line's intercepts
($a[0], $b[0]) = (-5..-1, 1..5)[NchooseK(10,2)];
# Get first line's slope
Context("Fraction");
$m[0] = Fraction(-$b[0]/$a[0]);
($n[0], $d[0]) = $m[0]->value;
# Gather integer points still in the viewing area
@pts = ();
for $i (-9..-1,1..9) {
  push(@pts, [$i, $m[0]*$i + $b[0]]) if ($n[0]*$i % $d[0] == 0
    && abs($m[0]*$i + $b[0]) <= 9
    && $i != $a[0]
  );
}
# Gather possible pairs [a,b] for the second line
@ab = ();
for $p (@pts) {
  for $b (-9..-1,1..9) {
    $m = Fraction(($p->[1] - $b)/($p->[0]));
    if ($m != 0) {
      my ($n, $d) = $m->value;
      push(@ab, [Real(-$b/$m)->value, $b]) if ($b % $n == 0
        && abs($b/$m) <= 9
        && $b != $b[0]
      );
    }
  }
}
$ab = list_random(@ab);
($a[1], $b[1]) = @{$ab};
@g = map{gcd($a[$_], $b[$_])} (0,1);
Context("ImplicitPlane");
Context()->noreduce('(-x)+y', '(-x)-y');
@f = map{ImplicitPlane("$b[$_]/$g[$_] x + $a[$_]/$g[$_] y =  $a[$_]*$b[$_]/$g[$_]")->reduce} (0,1);
$gt = GraphTool(
  "{line, solid, ($a[0], 0), (0, $b[0])}",
  "{line, solid, ($a[1], 0), (0, $b[1])}",
)->with(
  availableTools => [ "LineTool" ],
  showInStatic => 0,
  cmpOptions => {
    list_checker => sub {
      my ($correct, $student, $ans, $value) = @_;
      return 0 if $ans->{isPreview};
      my $score = 0; my @errors; my $count = 1;
      my ($p1x, $p1y) = $correct->[0]->extract(3)->value;
      my ($p2x, $p2y) = $correct->[0]->extract(4)->value;
      my ($p3x, $p3y) = $correct->[1]->extract(3)->value;
      my ($p4x, $p4y) = $correct->[1]->extract(4)->value;
      my @stdform = (
        $p1y - $p2y, $p2x - $p1x, $p1x * $p2y - $p2x * $p1y,
        $p3y - $p4y, $p4x - $p3x, $p3x * $p4y - $p4x * $p3y
      );
      my $pointOnLine = sub {
        my ($a, $b, $c, $point) = @_;  my ($x, $y) = $point->value;
        return $a * $x + $b * $y + $c == 0;
      };
      my @line = (0, 0);
      for (@$student) {
        my $nth = Value::List->NameForNumber($count++);
        if ($pointOnLine->(@stdform[0..2],$_->extract(3)) && $pointOnLine->(@stdform[0..2],$_->extract(4))) {
          $line[0] = 1; $score += 1; next
        }
        if ($pointOnLine->(@stdform[3..5],$_->extract(3)) && $pointOnLine->(@stdform[3..5],$_->extract(4))) {
          $line[1] = 1; $score += 1; next
        }
        push(@errors, "The $nth line graphed is incorrect.");
      }
      return (($line[0] && $line[1]) ? $score : 0, @errors);
    }
  }
);
Context("Point");
Context()->strings->remove("NONE");
Context()->strings->add(
  "no solutions" => {},
);
for (
  "no solution",
  "none",
  "there are no solutions",
  "the solution set is empty",
  "solution set is empty",
  "empty set",
  "∅",
  "{}",
  "{ }",
  "{  }",
  "{   }"
) {
  Context()->strings->add($_ => {alias => "no solutions"});
}
Context()->strings->add(
  "infinitely many solutions" => {},
);
for (
  "infinitely many",
  "infinite solution set",
  "infinite",
  "the solution set is infinite",
  "solution set is infinite"
) {
  Context()->strings->add($_ => {alias => "infinitely many solutions"});
}
$x = ($a[1]*$a[0]*$b[0] - $a[0]*$a[1]*$b[1])/($b[0]*$a[1] - $a[0]*$b[1]);
$y = (-$b[1]*$a[0]*$b[0] + $b[0]*$a[1]*$b[1])/($b[0]*$a[1] - $a[0]*$b[1]);
$solution = Point($x, $y)->cmp(
  mathQuillOpts => 'disabled'
);
@fplot = map{Formula("$b[$_] - $b[$_]/$a[$_] x")->reduce} (0,1);
@low = map{(10 + $b[$_])*$a[$_]/$b[$_]} (0,1);
@high = map{(-10 + $b[$_])*$a[$_]/$b[$_]} (0,1);
@dom0 = num_sort(-10, 10, $low[0], $high[0]);
($l[0], $r[0]) = (@dom0)[1,2];
@dom1 = num_sort(-10, 10, $low[1], $high[1]);
($l[1], $r[1]) = (@dom1)[1,2];

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      xmin=-10,
      xmax=10,
      ymin=-10,
      ymax=10,
      xtick = {-10,-5,5,10},
      minor xtick = {-10,...,10},
      ytick = {-10,-5,5,10},
      minor ytick = {-10,...,10},
      width=0.47\orccaprintwidth,
    ]
    \addplot+[domain=$l[0]:$r[0]]{$fplot[0]};
    \addplot+[domain=$l[1]:$r[1]]{$fplot[1]};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Solve the given system of linear equations graphically.

[`\newcommand{\amp}{&}
                  \left\{
                    \begin{alignedat}{3}
                      \amp [$f[0]]\\
                      \amp [$f[1]]
                    \end{alignedat}
                  \right.
                `]

[_]{$gt}{1}

[@KeyboardInstructions('Enter the solution here. If there are no solutions, enter [|no solutions|]*. If there are infinitely many solutions, enter [|infinitely many solutions|]*.')@]**

[_]{$solution}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
If we graph these lines, we find they cross.

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a coordinate plot of two lines that cross at ($x, $y)")@]*  <<

The only solution is where they cross, which is at [`[$solution->{rh_ans}{correct_ans}]\text{.}`]

END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
