#############################################
###       Generated from PreTeXt source      
###       on 2026-02-20T13:16:09-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.11.13)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "PCCmacros.pl",
  "parserGraphTool.pl",
  "Pd.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
$gt = GraphTool(
  "{line, solid, (2, 1), (3, 4)}",
  "{line, solid, (-1, -1), (2, -2)}"
)->with(
  availableTools => [ "LineTool" ],
  showInStatic => 0,
  cmpOptions => {
    list_checker => sub {
      my ($correct, $student, $ans, $value) = @_;
      return 0 if $ans->{isPreview};
      my $score = 0; my @errors; my $count = 1;
      my ($p1x, $p1y) = $correct->[0]->extract(3)->value;
      my ($p2x, $p2y) = $correct->[0]->extract(4)->value;
      my ($p3x, $p3y) = $correct->[1]->extract(3)->value;
      my ($p4x, $p4y) = $correct->[1]->extract(4)->value;
      my @stdform = (
        $p1y - $p2y, $p2x - $p1x, $p1x * $p2y - $p2x * $p1y,
        $p3y - $p4y, $p4x - $p3x, $p3x * $p4y - $p4x * $p3y
      );
      my $pointOnLine = sub {
        my ($a, $b, $c, $point) = @_;  my ($x, $y) = $point->value;
        return $a * $x + $b * $y + $c == 0;
      };
      my @line = (0, 0);
      for (@$student) {
        my $nth = Value::List->NameForNumber($count++);
        if ($pointOnLine->(@stdform[0..2],$_->extract(3)) && $pointOnLine->(@stdform[0..2],$_->extract(4))) {
          $line[0] = 1; $score += 1; next
        }
        if ($pointOnLine->(@stdform[3..5],$_->extract(3)) && $pointOnLine->(@stdform[3..5],$_->extract(4))) {
          $line[1] = 1; $score += 1; next
        }
        push(@errors, "The $nth line graphed is incorrect.");
      }
      return (($line[0] && $line[1]) ? $score : 0, @errors);
    }
  }
);
Context("Point");
$solution = Point(1,-2);

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      width=0.47\orccaprintwidth
    ]
    \addplot[firstcurve, domain=-0.66:4]{3*(x-2)+1} node[pos=1,sloped,above left] {\(y=3(x-2)+1\)};
    \addplot[secondcurve, domain=-7:7]{-1/2*(x+1)-1} node[pos=0,sloped,above right] {\(y=-\frac{1}{2}(x+1)-1\)};
    \addplot[guideline,->] coordinates {(2,1) (3,1) (3,4)};
    \addplot[guideline,->] coordinates {(2,1) (1,1) (1,-2)};
    \addplot[guideline,->] coordinates {(-1,-1) (1,-1) (1,-2)};
    \addplot[guideline,->] coordinates {(-1,-1) (-3,-1) (-3,0)};
    \addplot[soliddot,color=firstcolor] coordinates {(2,1) (3,4) (1,-2)};
    \addplot[soliddot,color=secondcolor] coordinates {(-1,-1) (1,-2) (-3,0)};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Solve the following system of equations by graphing. Note that the equations are in point-slope form. To review plotting a line equation that is in point-slope form, see Section[$NBSP]*10.

[```\newcommand{\amp}{&}
              \left\{
                \begin{alignedat}{3}
                  y \amp {}={} \amp 3(x-2) \amp {}+{} \amp 1 \\
                  y \amp {}={} \amp -\tfrac{1}{2}(x+1) \amp {}-{} \amp 1
                \end{alignedat}
              \right.
            ```]

[_]{$gt}{1}

[@KeyboardInstructions('Enter the solution here.')@]**

[_]{$solution}{5}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
Both of these equations are written in point-slope form. The first equation, [`y=3(x-2)+1\text{,}`] passes through [`(2,1)`] and has slope [`3\text{.}`] The second equation, [`y=-\frac{1}{2}(x+1)-1\text{,}`] passes through [`(-1,-1)`] and has slope [`-\frac{1}{2}\text{.}`] We can use this information to graph both lines.

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a Cartesian grid with two intersecting lines; one line passes through (2,1) with slope 3; the other line passes through (-1,-1) with slope -1/2")@]*  <<

It appears that the two lines intersect where [`x=1`] and [`y=-2\text{,}`] so the solution to the system of equations would be the point [`(1,-2)\text{.}`] However we should check that [`(1,-2)`] actually works as a solution to each of the original equations.

[```\newcommand{\substitute}[1]{{\color{blue}{{#1}}}}\newcommand{\confirm}[1]{\stackrel{\checkmark}{#1}}\newcommand{\wonder}[1]{\stackrel{\text{?}}{#1}}\newcommand{\amp}{&}
\begin{aligned}
y\amp=3(x-2)+1\amp y\amp=-\frac{1}{2}(x+1)-1\\
\substitute{-2}\amp\wonder{=}3(\substitute{1}-2)+1\amp \substitute{-1}\amp\wonder{=}-\frac{1}{2}(\substitute{1}+1)-1\\
-2\amp\confirm{=}3(-1)+1\amp -2\amp\confirm{=}-\frac{1}{2}(2)-1
\end{aligned}
```]

This verifies that [`(1,-2)`] is the solution.

END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
