#############################################
###       Generated from PreTeXt source      
###       on 2026-02-20T13:16:09-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.11.5.29)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "PCCmacros.pl",
  "PGchoicemacros.pl",
  "parserImplicitPlane.pl",
  "contextForm.pl",
  "contextFraction.pl",
  "parserGraphTool.pl",
  "Pd.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
# Get first line's intercepts
($a[0], $b[0]) = (-5..-1, 1..5)[NchooseK(10,2)];
# Get first line's slope
Context("Fraction");
Context()->noreduce('(-x)+y', '(-x)-y');
$m[0] = Fraction(-$b[0]/$a[0]);
($n[0], $d[0]) = $m[0]->value;
# Gather integer points still in the viewing area
@pts = ();
for $i (-9..-1,1..9) {
  push(@pts, [$i, $m[0]*$i + $b[0]]) if ($n[0]*$i % $d[0] == 0
    && abs($m[0]*$i + $b[0]) <= 9
    && $i != $a[0]
  );
}
$pt = list_random(@pts);
($ptx, $pty) = @$pt;
$b[1] = list_random(-9..$b[0]-1,$b[0]+1..9);
$m[1] = Fraction(($pt->[1] - $b[1])/($pt->[0]));
$f[1] = Formula("$m[1] x + $b[1]")->reduce;
$g[0] = gcd($a[0], $b[0]);
Context("ImplicitPlane");
Context()->noreduce('(-x)+y', '(-x)-y');
$f[0] = ImplicitPlane("$b[0]/$g[0] x + $a[0]/$g[0] y =  $a[$0]*$b[$0]/$g[$0]")->reduce;
$gt = GraphTool(
  "{line, solid, ($a[0], 0), (0, $b[0])}",
  "{line, solid, ($ptx, $pty), (0, $b[1])}",
)->with(
  availableTools => [ "LineTool" ],
  showInStatic => 0,
  cmpOptions => {
    list_checker => sub {
      my ($correct, $student, $ans, $value) = @_;
      return 0 if $ans->{isPreview};
      my $score = 0; my @errors; my $count = 1;
      my ($p1x, $p1y) = $correct->[0]->extract(3)->value;
      my ($p2x, $p2y) = $correct->[0]->extract(4)->value;
      my ($p3x, $p3y) = $correct->[1]->extract(3)->value;
      my ($p4x, $p4y) = $correct->[1]->extract(4)->value;
      my @stdform = (
        $p1y - $p2y, $p2x - $p1x, $p1x * $p2y - $p2x * $p1y,
        $p3y - $p4y, $p4x - $p3x, $p3x * $p4y - $p4x * $p3y
      );
      my $pointOnLine = sub {
        my ($a, $b, $c, $point) = @_;  my ($x, $y) = $point->value;
        return $a * $x + $b * $y + $c == 0;
      };
      my @line = (0, 0);
      for (@$student) {
        my $nth = Value::List->NameForNumber($count++);
        if ($pointOnLine->(@stdform[0..2],$_->extract(3)) && $pointOnLine->(@stdform[0..2],$_->extract(4))) {
          $line[0] = 1; $score += 1; next
        }
        if ($pointOnLine->(@stdform[3..5],$_->extract(3)) && $pointOnLine->(@stdform[3..5],$_->extract(4))) {
          $line[1] = 1; $score += 1; next
        }
        push(@errors, "The $nth line graphed is incorrect.");
      }
      return (($line[0] && $line[1]) ? $score : 0, @errors);
    }
  }
);
Context("Point");
Context()->strings->remove("NONE");
Context()->strings->add(
  "no solutions" => {},
);
for (
  "no solution",
  "none",
  "there are no solutions",
  "the solution set is empty",
  "solution set is empty",
  "empty set",
  "∅",
  "{}",
  "{ }",
  "{  }",
  "{   }"
) {
  Context()->strings->add($_ => {alias => "no solutions"});
}
Context()->strings->add(
  "infinitely many solutions" => {},
);
for (
  "infinitely many",
  "infinite solution set",
  "infinite",
  "the solution set is infinite",
  "solution set is infinite"
) {
  Context()->strings->add($_ => {alias => "infinitely many solutions"});
}
$solution = Point(@$pt)->cmp(
  mathQuillOpts => 'disabled'
);
$fplot[0] = Formula("$b[0] - $b[0]/$a[0] x")->reduce;

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      xmin=-10,
      xmax=10,
      ymin=-10,
      ymax=10,
      xtick = {-10,-5,5,10},
      minor xtick = {-10,...,10},
      ytick = {-10,-5,5,10},
      minor ytick = {-10,...,10},
      width=0.47\orccaprintwidth,
    ]
    \addplot[firstcurve,-,domain=-10:10]{$fplot[0]};
    \addplot[secondcurve,-,domain=-10:10]{$f[1]};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Solve the given system of linear equations graphically.

[`\newcommand{\amp}{&}
                  \left\{
                    \begin{alignedat}{3}
                      \amp y=[$f[1]]\\
                      \amp [$f[0]]
                    \end{alignedat}
                  \right.
                `]

[_]{$gt}{1}

[@KeyboardInstructions('Enter the solution here. If there are no solutions, enter [|no solutions|]*. If there are infinitely many solutions, enter [|infinitely many solutions|]*.')@]**

[_]{$solution}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
If we graph these lines, we find they are parallel and never cross.

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a coordinate plot of two parallel lines")@]*  <<

So the there are no solutions.

END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
