#############################################
###       Generated from PreTeXt source      
###       on 2026-02-20T13:16:09-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.11.5.42)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "scaffold.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "contextCurrency.pl",
  "contextForm.pl",
  "parserLinearRelation.pl",
  "parserGraphTool.pl",
  "randomPerson.pl",
  "Pd.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
$nx = list_random([2,10], [2,30], [2,40], [3,15], [3,30], [3,45], [4,20], [4,40]);
($n,$x) = @$nx;
$y = random(2,5);
$t = $x + $n*$y;
$d = $x - $y;
Context("Currency")->flags->set(trimTrailingZeros => 1);
$total = Currency($t);
$difference = Currency($d);
$hammer = Currency($x);
$nails = Currency($y);
@f = (Formula("($t-x)/$n"), Formula("x-$d"));
Context("LinearRelation")->variables->are(x=>'Real',y=>'Real');
@eq=(LinearRelation("x + $n y = $t"),LinearRelation("x = y + $d"));
$m=MultiAnswer(@eq)->with(singleResult=>0,
checker=>sub{
my($c,$s,$self)=@_;
my @s=@{$s};
my @c=@{$c};
if (($c[0]==$s[0]&&$c[1]==$s[1])||($c[0]==$s[1]&&$c[1]==$s[0])){return [1,1]}
elsif ($c[0]==$s[0]||$c[1]==$s[0]){return [1,0]}
elsif ($c[0]==$s[1]||$c[1]==$s[1]){return [0,1]}
else {return [0,0]}
});
$ymax=int(($t/$n)/10+0.999)*10;
$ymin=int(-$d/10-0.999)*10;
$xmax=int($t/10+0.999)*10;
$gt=GraphTool("{line,solid,(0,$t/$n),($t,0)}","{line,solid,(0,-$d),($d,0)}")->with(
availableTools=>["LineTool"],showInStatic=>0,bBox=>[-5,$ymax,$xmax,$ymin],
ticksDistanceX=>5,ticksDistanceY=>5,minorTicksX=>4,minorTicksY=>4,
cmpOptions=>{list_checker=>sub{
my($c,$s,$a,$v)=@_;
return 0 if $ans->{isPreview};
my $score=0; my @errors; my $count=1;
my ($x1,$y1)=$c->[0]->extract(3)->value;
my ($x2,$y2)=$c->[0]->extract(4)->value;
my ($x3,$y3)=$c->[1]->extract(3)->value;
my ($x4,$y4)=$c->[1]->extract(4)->value;
my @std=($y1-$y2,$x2-$x1,$x1*$y2-$x2*$y1,$y3-$y4,$x4-$x3,$x3*$y4-$x4*$y3);
my $p=sub{my ($A,$B,$C,$pt)=@_; my ($x,$y)=$pt->value; return $A*$x+$B*$y+$C==0;};
my @l=(0,0);
for (@$s) {
my $nth=Value::List->NameForNumber($count++);
if ($p->(@std[0..2],$_->extract(3))&&$p->(@std[0..2],$_->extract(4))) {$l[0]=1;$score += 1; next}
if ($p->(@std[3..5],$_->extract(3))&&$p->(@std[3..5],$_->extract(4))) {$l[1]=1;$score += 1; next}
push(@errors, "The $nth line graphed is incorrect.");
}
return (($l[0]&&$l[1]) ? $score : 0,@errors);
}});
$p=randomPerson();
$w = numberWord($n);

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
\begin{axis}[xmin=-5,xmax=$xmax,ymin=$ymin,ymax=$ymax,xtick={-100,-95,...,100},ytick={-100,-95,...,100},minor xtick={-100,-99,...,100},minor ytick={-100,-99,...,100}]
\addplot[firstcurve,-,domain=-5:$xmax]{$f[0]};
\addplot[secondcurve,-,domain=-5:$xmax]{$f[1]};
\end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
At the local hardware store, [$p] bought a hammer and [$w] boxes of nails. The total cost was [`[$total]\text{.}`] The hammer costs [`[$difference]`] more than a box of nails. How much does a hammer cost and how much does one box of nails cost?

END_PGML

############################################################
# Scaffold
############################################################
Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Write two equations that form a system for this scenario.

[@KeyboardInstructions('Use [|x|]* as the variable for the cost of the hammer, and [|y|]* for the variable representing the cost of one box of nails.')@]**

[_]{$m}{16}

[_]{$m}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
Since the total cost of the hammer and [$w] boxes of nails is [`[$total]\text{,}`] one equation is [`[$eq[0]]\text{.}`]

Since the hammer costs [`[$difference]`] more than a box of nails, another equation is [`[$eq[1]]\text{.}`]

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Plot the lines for the system of two equations.

[_]{$gt}{1}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
Plotting the two slope-intercept equations:

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a Cartesian grid with two intersecting lines; one line passes through (0,$t/$n) and ($t,0); the other line passes through (0,-$d) and ($d,0); the lines cross at ($x,$y)")@]*  <<

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Based on the graph, how much does the hammer cost?

[_]{$hammer}{16}

How much does a box of nails cost?

[_]{$nails}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The lines cross at [`([$x],[$y])\text{.}`] So a hammer costs [`[$hammer]\text{,}`] and a box of nails costs  [$nails].

END_PGML_SOLUTION

Section::End();

Scaffold::End();

############################################################
# End Problem
############################################################

ENDDOCUMENT();
