#############################################
###       Generated from PreTeXt source      
###       on 2026-02-20T13:16:09-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.11.5.41)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "scaffold.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "contextCurrency.pl",
  "contextForm.pl",
  "parserLinearRelation.pl",
  "parserGraphTool.pl",
  "randomPerson.pl",
  "Pd.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
$p2 = random(1,2,0.25);
$p1 = $p2 + random(0.5,1,0.25);
$a = 4*$p2/(gcd(4*$p1,4*$p2));
$b = 4*$p1/(gcd(4*$p1,4*$p2));
$x = random(3*$a,6*$a,$a);
$y = random(3*$b,6*$b,$b);
$t = $p1*$x + $p2*$y;
$count = ($x + $y)/2;
@f = (Formula("($t-$p1*x)/$p2"), Formula("2*$count - x"));
Context("Currency")->flags->set(trimTrailingZeros => 1);
$total = Currency($t);
$price1 = Currency($p1);
$price2 = Currency($p2);
Context("LinearRelation")->variables->are(x=>'Real',y=>'Real');
@eq=(LinearRelation("$p1 x + $p2 y = $t"),LinearRelation("x + y = 2*$count"));
$m=MultiAnswer(@eq)->with(singleResult=>0,
checker=>sub{
my($c,$s,$self)=@_;
my @s=@{$s};
my @c=@{$c};
if (($c[0]==$s[0]&&$c[1]==$s[1])||($c[0]==$s[1]&&$c[1]==$s[0])){return [1,1]}
elsif ($c[0]==$s[0]||$c[1]==$s[0]){return [1,0]}
elsif ($c[0]==$s[1]||$c[1]==$s[1]){return [0,1]}
else {return [0,0]}
});
$ymax=int(max(2*$count, $t/$p2)/10+0.999)*10;
$xmax=int(max(2*$count, $t/$p1)/10+0.999)*10;
$gt=GraphTool("{line,solid,(0,$t/$p2),($t/$p1,0)}","{line,solid,(0,2*$count),(2*$count,0)}")->with(
availableTools=>["LineTool"],showInStatic=>0,bBox=>[-5,$ymax,$xmax,-5],
ticksDistanceX=>5,ticksDistanceY=>5,minorTicksX=>4,minorTicksY=>4,
cmpOptions=>{list_checker=>sub{
my($c,$s,$a,$v)=@_;
return 0 if $ans->{isPreview};
my $score=0; my @errors; my $count=1;
my ($x1,$y1)=$c->[0]->extract(3)->value;
my ($x2,$y2)=$c->[0]->extract(4)->value;
my ($x3,$y3)=$c->[1]->extract(3)->value;
my ($x4,$y4)=$c->[1]->extract(4)->value;
my @std=($y1-$y2,$x2-$x1,$x1*$y2-$x2*$y1,$y3-$y4,$x4-$x3,$x3*$y4-$x4*$y3);
my $p=sub{my ($A,$B,$C,$pt)=@_; my ($x,$y)=$pt->value; return $A*$x+$B*$y+$C==0;};
my @l=(0,0);
for (@$s) {
my $nth=Value::List->NameForNumber($count++);
if ($p->(@std[0..2],$_->extract(3))&&$p->(@std[0..2],$_->extract(4))) {$l[0]=1;$score += 1; next}
if ($p->(@std[3..5],$_->extract(3))&&$p->(@std[3..5],$_->extract(4))) {$l[1]=1;$score += 1; next}
push(@errors, "The $nth line graphed is incorrect.");
}
return (($l[0]&&$l[1]) ? $score : 0,@errors);
}});
$p=randomPerson();

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
\begin{axis}[xmin=-5,xmax=$xmax,ymin=-5,ymax=$ymax,xtick={-100,-95,...,100},ytick={-100,-95,...,100},minor xtick={-100,-99,...,100},minor ytick={-100,-99,...,100}]
\addplot[firstcurve,-,domain=-5:$xmax]{$f[0]};
\addplot[secondcurve,-,domain=-5:$xmax]{$f[1]};
\end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
[$p] is organizing an office lunch party. [$p->Their] budget for beverages is [`[$total]`] and [$p->they] will try to spend all of it. [$p->They] [$p->verb('assume')] that [`[$count]`] people will attend, and there should be two beverages available per person. [$p] has decided to order cans of a fancy beverage that cost [`[$price1]`] each, and cans of a cheaper beverage that cost [`[$price2]`]each. How many of each should [$p->they] order?

END_PGML

############################################################
# Scaffold
############################################################
Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Write two equations that form a system for this scenario.

[@KeyboardInstructions('Use [|x|]* as the variable representing how many cans of the more expensive beverage to buy, and [|y|]* for the variable representing how many cans of the less expensive beverage to buy.')@]**

[_]{$m}{16}

[_]{$m}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
Since there are [`[$count]`] people and there should be two drinks per person, one equation is [`[$eq[0]]\text{.}`]

Since the first type costs [`[$price1]`] each, the second type costs [`[$price2]`] each, and the total should be [`[$total]\text{,}`] another equation is [`[$eq[1]]\text{.}`]

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Plot the lines for the system of two equations.

[_]{$gt}{1}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
Plotting the two slope-intercept equations:

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a Cartesian grid with two intersecting lines; one line passes through (0,$t/$p2) and ($t/$p1,0); the other line passes through (0,2*$count) and (2*$count,0); the lines cross at ($x,$y)")@]*  <<

END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Based on the graph, how many of the more expensive beverage should [$p] buy?

[_]{$x}{16}

How many of the less expensive beverage should [$p] buy?

[_]{$y}{16}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The lines cross at [`([$x],[$y])\text{.}`] So [$p] should buy [$x] of the more expensive drink, and [$y] of the less expensive drink.

END_PGML_SOLUTION

Section::End();

Scaffold::End();

############################################################
# End Problem
############################################################

ENDDOCUMENT();
