#############################################
###       Generated from PreTeXt source      
###       on 2026-02-21T15:38:35-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.12.5.53)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "parserNumberWithUnits.pl",
  "parserLinearRelation.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
$n = random(2,8);
$w = numberWord($n);
$d = random(5,10);
$diff = NumberWithUnits("$d ft");
$p = random(100,200);
$perimeter = NumberWithUnits("$p ft");
Context("LinearRelation")->variables->are(L => 'Real', W => 'Real');
@eq = (LinearRelation("L = $n W - $d"), LinearRelation("2L + 2W = $p"));
$m = MultiAnswer(@eq)->with(
  singleResult => 0,
  checker => sub {
      my ( $c, $s, $self ) = @_;
      return [1,1] if ($c->[0] == $s->[0] && $c->[1] == $s->[1] || $c->[1] == $s->[0] && $c->[0] == $s->[1]);
      return [1,0] if ($c->[0] == $s->[0] || $c->[1] == $s->[0]);
      return [0,1] if ($c->[0] == $s->[1] || $c->[1] == $s->[1]);
      return [0,0];
  }
);
$W = NumberWithUnits(($p+2*$d)/(2*$n+2), "ft");
$L = NumberWithUnits($n*($p+2*$d)/(2*$n+2) - $d, "ft");

############################################################
# Body
############################################################

BEGIN_PGML
A rectangle's length is [`[$diff]`] shorter than [$w] times its width. The rectangle's perimeter is [`[$perimeter]\text{.}`] Find the rectangle's length and width.

[@KeyboardInstructions('Using [`L`] and [`W`] as the variables, give two equations that form a system of equations that is appropriate for this exercise.')@]**

[_]{$m}{16}

[_]{$m}{16}

[@KeyboardInstructions(q"Solve the system and enter the rectangle's length here.")@]**

[_]{$L}{16}

[@KeyboardInstructions('Enter the width here.')@]**

[_]{$W}{16}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
