#############################################
###       Generated from PreTeXt source      
###       on 2026-02-24T13:47:08-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.4.8.102)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "contextPercent.pl",
  "parserLinearRelation.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
$group = list_random('Democrats', 'Republicans', 'retirees', 'married');
$a = random(1000, 2000, 100);
$b = random(2100, 4000, 100);
$total = $a + $b;
Context("Percent");
$p1 = random(0.40, 0.48, 0.01);
$P1 = Percent($p1);
$p2 = random(0.49, 0.55, 0.01);
$P2 = Percent($p2);
$pfinal = Round(($p1*$a + $p2*$b)/($a + $b),3);
$Pfinal = Percent($pfinal);
Context("LinearRelation")->variables->are(A => 'Real', B => 'Real');
@eq = (LinearRelation("A + B = $total"), LinearRelation("$p1 A + $p2 B = $pfinal*$total"));
$m = MultiAnswer(@eq)->with(
  singleResult => 0,
  checker => sub {
      my ( $c, $s, $self ) = @_;
      return [1,1] if ($c->[0] == $s->[0] && $c->[1] == $s->[1] || $c->[1] == $s->[0] && $c->[0] == $s->[1]);
      return [1,0] if ($c->[0] == $s->[0] || $c->[1] == $s->[0]);
      return [0,1] if ($c->[0] == $s->[1] || $c->[1] == $s->[1]);
      return [0,0];
  }
);
$B = Round(($p1*$total - $pfinal*$total)/($p1 - $p2));
$A = $total - $B;

############################################################
# Body
############################################################

BEGIN_PGML
Asterton and Balsamburg were two towns located close to each other, and recently they merged into one city. Asterton had a population with [`[$P1]`] [$group]. Balsamburg had a population with [`[$P2]`] [$group]. After the merge, the new city has a total of [`[$total]`] residents, with [`[$Pfinal]`] [$group]. How many residents did each town have before the merge?

[@KeyboardInstructions('Using [`A`] and [`B`] as the variables, give two equations that form a system of equations that is appropriate for this exercise.')@]**

[_]{$m}{16}

[_]{$m}{16}

[@KeyboardInstructions('Solve the system and enter the population of Asterton here.')@]**

[_]{$A}{16}

[@KeyboardInstructions('Enter the population of Balsamburg here.')@]**

[_]{$B}{16}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
