#############################################
###       Generated from PreTeXt source      
###       on 2026-02-24T13:47:08-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.4.8.100)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "contextCurrency.pl",
  "contextPercent.pl",
  "parserLinearRelation.pl",
  "randomPerson.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
$p = randomPerson();
Context("Percent");
$r1 = random(0.01, 0.03, 0.005);
$rate1 = Percent($r1);
$r2 = random(0.035, 0.07, 0.005);
$rate2 = Percent($r2);
Context("Currency")->flags->set(trimTrailingZeros => 1);
$x = random(1000, 2000, 100);
$y = random(2100, 4000, 100);
$X = Currency($x);
$Y = Currency($y);
$t = $x + $y;
$total = Currency($t);
$i = $r1 * $x + $r2 * $y;
$interest = Currency($i);
Context("LinearRelation")->variables->are(x => 'Real', y => 'Real');
@eq = (LinearRelation("x + y = $t"), LinearRelation("$r1 x + $r2 y = $i"));
$m = MultiAnswer(@eq)->with(
  singleResult => 0,
  checker => sub {
      my ( $c, $s, $self ) = @_;
      return [1,1] if ($c->[0] == $s->[0] && $c->[1] == $s->[1] || $c->[1] == $s->[0] && $c->[0] == $s->[1]);
      return [1,0] if ($c->[0] == $s->[0] || $c->[1] == $s->[0]);
      return [0,1] if ($c->[0] == $s->[1] || $c->[1] == $s->[1]);
      return [0,0];
  }
);

############################################################
# Body
############################################################

BEGIN_PGML
[$p] invested a total of [`[$total]`] in two investments. [$p->Their] savings account pays [`[$rate1]`] interest annually. A riskier stock investment earned [`[$rate2]`] at the end of the year. At the end of the year, [$p] earned a total of [`[$interest]`] in interest. How much money did [$p->they] invest in each account?

[@KeyboardInstructions('Using [`x`] (savings investment) and [`y`] (stock investment) as the variables, give two equations that form a system of equations that is appropriate for this exercise.')@]**

[_]{$m}{16}

[_]{$m}{16}

[@KeyboardInstructions('Solve the system and enter the amount invested in savings here.')@]**

[_]{$X}{16}

[@KeyboardInstructions('Enter the amount invested in stock here.')@]**

[_]{$Y}{16}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
