#############################################
###       Generated from PreTeXt source      
###       on 2026-02-12T21:52:51-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(2.1.9)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "parserGraphTool.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context('Numeric');
@xy = random_subset(4, 100..200);
@x = ((@xy)[0,1]);
@y = ((@xy)[2,3]);
@p = map {Point($x[$_], $y[$_])} (0,1);
$gt = GraphTool("{point,$p[0]}", "{point,$p[1]}")->with(
  availableTools => ['PointTool'],
  bBox => [-30,220,220,-30],
  gridX => 10,
  gridY => 10,
  ticksDistanceX => 50,
  ticksDistanceY => 50,
  snapSizeX => 5,
  snapSizeY => 5,
  showInStatic => 0,
  showCoordinateHints => 0,
  cmpOptions => {
    list_checker => sub {
      my ($c, $s, $ans, $value) = @_;
      return 0 if $ans->{isPreview};
      my $score = 0;
      my @errors;
      my ($p, $q) = (0, 0);
      for (0 .. $#$s) {
        if (!$p && (abs(($s->[$_]->extract(2)->value)[0] - $x[0]) < 5 && abs(($s->[$_]->extract(2)->value)[1] - $y[0]) < 5)) { $p = 1; ++$score; next; }
        if (!$q && (abs(($s->[$_]->extract(2)->value)[0] - $x[1]) < 5 && abs(($s->[$_]->extract(2)->value)[1] - $y[1]) < 5)) { $q = 1; ++$score; next; }
        my $nth = Value::List->NameForNumber($_ + 1);
        push(@errors, "The $nth point graphed is incorrect.");
      }   
      return ($score, @errors);
    }
  }
);

############################################################
# Body
############################################################

BEGIN_PGML
Make a Cartesian plot with the indicated points marked.

[`[$p[0]], [$p[1]]`]

[_]{$gt}{1}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
