#############################################
###       Generated from PreTeXt source      
###       on 2026-02-12T21:52:51-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(2.4.5.37)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "contextFraction.pl",
  "parserGraphTool.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
($rise, $run) = random_coprime([1..5],[2..5]);
($h,$k) = random_subset(2, -5..-1,1..3);
($u,$v) = ($h+$run,$k+$rise);
Context("Numeric");
$A = Point($h,$k);
$B = Point($u,$v);
Context("Fraction");
$m = Fraction($rise,$run);              
$gt = GraphTool("{line,solid,$A,$B}")->with(
showInStatic => 0,
availableTools => ["LineTool"],
showCoordinateHints => 0,
bBox => [-8, 8, 8, -8],
cmpOptions => {
list_checker => sub {
my ($correct, $student, $ans, $value) = @_;
return 0 if $ans->{isPreview};
my @errors;
my $count = 0;
my $distinct = 0;
my $extra = 0;
STUDENTS:
for my $j (@$student) {
my $nth = Value::List->NameForNumber(++$count);
my ($p1x, $p1y) = $j->extract(3)->value;
my ($p2x, $p2y) = $j->extract(4)->value;
if ($p1x == $p2x) {
push(@errors, "The $nth plotted line is a vertical line, and therefore does not have slope $m.");
$extra++;
next STUDENTS;
}
my $context = Context();
Context("Fraction");
my $slope = Fraction($p2y - $p1y, $p2x - $p1x);
Context($context);
if ($slope != $m) {
push(@errors, "The $nth plotted line has slope $slope, not slope $m.");
$extra++;
next STUDENTS;
}
if ($k != $slope*($h - $p1x) + $p1y) {
push(@errors, "The $nth plotted line does not pass through $A.");
$extra++;
next STUDENTS;
}
for my $i (0..($count-2)) {
my ($x,$y) = $student->[$i]->extract(3)->value;
next STUDENTS if ($y == $m*($x - $p1x) + $p1y);
}
$distinct++;
}
return (max(scalar(@{$correct}),scalar(@{$student}))*min($distinct,1)/1/(($extra)?2:1), @errors);
}
}
);

############################################################
# Body
############################################################

BEGIN_PGML
Plot a line through the given point that has the given slope.

Through [`[$A]`] with slope [`[$m]\text{.}`]

[_]{$gt}{1}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
