#############################################
###       Generated from PreTeXt source      
###       on 2026-02-12T21:52:51-08:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(2.4.5.43)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "contextFraction.pl",
  "parserGraphTool.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
($rise, $run) = random_coprime([-5..-1],[2..5]);
Context("Numeric");
@A = ();
for my $x (-8..(8-$run)) {
for my $y (-8..(-9+$rise)) {
push(@A, Point($x,$y));
}
}
for my $x (-8..(-9+$run)) {
for my $y ((-8+$rise)..(8-$rise)) {
push(@A, Point($x,$y));
}
}
%B = map{$_ => $_+Point($run,$rise)}(@A);
@gtlines = map{'{line,solid,' . $_ . ',' . $B{$_} .'}'} (@A);
Context("Fraction");
$m = Fraction($rise,$run);
$gt = GraphTool(join(',',@gtlines))->with(
showInStatic => 0,
availableTools => ["LineTool"],
bBox => [-8, 8, 8, -8],
cmpOptions => {
list_checker => sub {
my ($correct, $student, $ans, $value) = @_;
return 0 if $ans->{isPreview};
my @errors;
my $count = 0;
my $extra = 0;
my $distinct = 0;
STUDENTS:
for my $j (@$student) {
my $nth = Value::List->NameForNumber(++$count);
my ($p1x, $p1y) = $j->extract(3)->value;
my ($p2x, $p2y) = $j->extract(4)->value;
if ($p1x == $p2x) {
push(@errors, "The $nth plotted line is a vertical line, and therefore does not have slope $m.");
$extra++;
next STUDENTS;
}
my $context = Context();
Context("Fraction");
my $slope = Fraction($p2y - $p1y, $p2x - $p1x);
Context($context);
if ($slope != $m) {
push(@errors, "The $nth plotted line has slope $slope, not slope $m.");
$extra++;
next STUDENTS;
}
for my $i (0..($count-2)) {
my ($x,$y) = $student->[$i]->extract(3)->value;
next STUDENTS if ($y == $m*($x - $p1x) + $p1y);
}
$distinct++;
}
push(@errors, "You've only plotted $distinct correct line".(($distinct > 1)?'s':'').".") if ($distinct > 0 && $distinct < 3);
return (scalar(@{$correct})*min($distinct,3)/3/(($extra)?2:1), @errors);
}
}
);

############################################################
# Body
############################################################

BEGIN_PGML
Plot at least three lines that each have slope the given slope.

[`[$m]`]

[_]{$gt}{1}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
