#############################################
###       Generated from PreTeXt source      
###       on 2025-08-09T22:53:06-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(2.4.4.41)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "contextForm.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context('Numeric');
$a = random(2,20,2);
$b = $a-1;
$answer = Formula("$a*(x+4)");
$answer2 = Formula("$b*x+($a*4+2)");
$evaluator2 = $answer2->cmp(
  checker => sub {
    my ($correct,$student,$ansHash) = @_;
    Value::Error('This is not a linear equation.') unless ($student->D('x')->D('x') == 0 );
    Value::Error('x=2 is not a solution to this equation.') unless ($student->eval(x=>2) == $a*6);
    Value::Error('This equation has infinitely many solutions, not just 2.') unless ($student != $answer);
    return 1;
  }
);

############################################################
# Body
############################################################

BEGIN_PGML
Fill in the right side of the equation to create a linear equation that meets the description.
a.  Create a linear equation with solution set [`\{2\}\text{.}`]

    [` [$a](x + 4) = `] [_]{$evaluator2}{20}

a.  Create a linear equation with _infinitely many solutions_.

    [` [$a](x + 4) = `] [_]{$answer}{20}

END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
