#############################################
###       Generated from PreTeXt source      
###       on 2025-08-09T22:53:07-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(3.5.14)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "bizarroArithmetic.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "parserAssignment.pl",
  "contextForm.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context()->operators->set(
  '+'=>{class=>'bizarro::BOP::add',isCommand=>1},
  '-'=>{class=>'bizarro::BOP::subtract',isCommand=>1},
);
parser::Assignment->Allow;
Context()->variables->add(y=>'Real');
$eq=Formula("y=-1/2x+6");
$ev=$eq->cmp(checker=>sub{
  my($co,$st,$aH)=@_;
  return 0 if$aH->{isPreview}||$co!=$st;
  $st=$aH->{student_formula};
  $co=$co->{original_formula} if defined $co->{original_formula};
  $st=Formula("$st");$co=Formula("$co");
  return 0 unless($co==$st);
  Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
  delete $co->{test_values},$st->{test_values};
  Value->Error("Your answer is an equation for this line, but it is not in simplified slope-intercept form")unless(($co==$st)or($st==$co));
  Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
  return 1;
})->withPostFilter(sub{my$ans=shift;
  $ans->{ans_message}="An equation for a line in slope-intercept form should start with 'y='"
    if $ans->{ans_message} eq "The left side of an assignment must be a variable or function";
  return $ans;
});

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      xmin=-2,
      xmax=8=,
      ymin=-2,
      ymax=8,
      ytick=,
      xtick=,
      width=0.47\orccaprintwidth
    ]
    \addplot+[domain=-2:8] {-1/2*x + 6};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

$image_2 = createLaTeXImage();
$image_2->addToPreamble(latexImagePreamble());
$image_2->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      xmin=-2,
      xmax=8,
      ymin=-2,
      ymax=8,
      ytick=,
      xtick=,
      width=0.47\orccaprintwidth
    ]
    \addplot+[domain=-2:8] {-1/2*x + 6};
    \addplot[soliddot] coordinates {(0,6)} node[above right] {\((0,6)\)};
    \addplot[soliddot] coordinates {(4,4)} node[below left] {\((4,4)\)};
    \addplot[guideline] coordinates {(0,6) (4,6)} node[above left] {run \(4\)};
    \addplot[guideline] coordinates {(4,6) (4,4)} node[pos=0.5, right] {rise \(-2\)};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Use the graph to write an equation of the line in slope-intercept form.

>> [@image(insertGraph($image_1), width=>282)@]*  <<

[_]{$ev}{20}

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
On the line, pick two points with easy-to-read integer coordinates so that we can calculate slope. It doesn't matter which two points we use; the slope will be the same.

>> [@image(insertGraph($image_2), width=>282)@]*  <<

With the slope and [`y`]-intercept found, we can write the line's equation:

[```
              y=-\frac{1}{2}x+6
            \text{.}```]

END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
