#############################################
###       Generated from PreTeXt source      
###       on 2025-08-09T22:53:07-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(3.8.11)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "niceTables.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "parserMultiAnswer.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());

############################################################
# PG Setup Code
############################################################
Context('Numeric');
@x=(-2..2);
@y=(-4.25)x5;
$pts=MultiAnswer(map{($_%2==0)?$x[$_/2]:$y[($_-1)/2]}(0..9))->with(
  singleResult => 1,
  allowBlankAnswers => 1,
  checker => sub {
    my ($correct,$student,$self) = @_;
    my @sx = @{$student}[0,2,4,6,8];
    my @sy = @{$student}[1,3,5,7,9];
    my $OK=1;
    if ($sy[0]!=-4.25) {
        $self->setMessage(1,"This point is not on the line.");
        $self->setMessage(2,"This point is not on the line.");
        $OK=0;
      }
    for my $i(1..4){
      if ($sy[$i]!=-4.25) {
        $self->setMessage(2*$i+1,"This point is not on the line.");
        $self->setMessage(2*$i+2,"This point is not on the line.");
        $OK=0;
      }
      for my $j(0..$i-1){
        if ($sy[$i]==-4.25 and $sx[$i]==$sx[$j]) {
          $self->setMessage(2*$i+1,"This point was already used.");
          $self->setMessage(2*$i+2,"This point was already used.");
          $OK=0;
        }
      }
    }
    return $OK;
  }
);

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      width=0.47\orccaprintwidth,
    ]
    \addplot+ {-4.25};
    \addplot[soliddot] coordinates {(-2,-4.25) (-1,-4.25) (0,-4.25) (1,-4.25) (2,-4.25)};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Suppose you need to plot the equation [`y=-4.25\text{.}`] Since the equation is in "[`y=`]" form, you decide to make a table of points. Fill out some points for this table.

>> [@DataTable(
  [
    [[PGML('[`x`]'),
      rowbottom => '2',
      midrule => '1',],
     PGML('[`y`]'),
    ],
    [PGML('[_]{$pts}{10}'),
     PGML('[_]{$pts}{10}'),
    ],
    [PGML('[_]{$pts}{10}'),
     PGML('[_]{$pts}{10}'),
    ],
    [PGML('[_]{$pts}{10}'),
     PGML('[_]{$pts}{10}'),
    ],
    [PGML('[_]{$pts}{10}'),
     PGML('[_]{$pts}{10}'),
    ],
    [PGML('[_]{$pts}{10}'),
     PGML('[_]{$pts}{10}'),
    ],
  ],
      align => 'cc',
      valign => 'middle',);@]* <<

END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
We can use whatever values for [`x`] that we like, as long as they are all different. The equation tells us the [`y`]-value has to be [`-4.25`] each time.

>> [@DataTable(
  [
    [[PGML('[`x`]'),
      rowbottom => '2',
      midrule => '1',],
     PGML('[`y`]'),
    ],
    [PGML('[`-2`]'),
     PGML('[`-4.25`]'),
    ],
    [PGML('[`-1`]'),
     PGML('[`-4.25`]'),
    ],
    [PGML('[`0`]'),
     PGML('[`-4.25`]'),
    ],
    [PGML('[`1`]'),
     PGML('[`-4.25`]'),
    ],
    [PGML('[`2`]'),
     PGML('[`-4.25`]'),
    ],
  ],
      align => 'cc',
      valign => 'middle',);@]* <<

Now that we have a table, we could use its values to assist with plotting the line.

>> [@image(insertGraph($image_1), width=>282)@]*  <<

END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
